<a href="https://colab.research.google.com/github/chrisbouton/MachineLearningMines/blob/main/1a_Getting_Started.ipynb" target="_parent"><img src="https://colab.research.google.com/assets/colab-badge.svg" alt="Open In Colab"/></a>

## CSCI 470 Activities and Case Studies

1. For all activities, you are allowed to collaborate with a partner. 
1. For case studies, you should work individually and are **not** allowed to collaborate.

By filling out this notebook and submitting it, you acknowledge that you are aware of the above policies and are agreeing to comply with them.

Some considerations with regard to how these notebooks will be graded:

1. Cells in which "# YOUR CODE HERE" is found are the cells where your graded code should be written.
2. In order to test out or debug your code you may also create notebook cells or edit existing notebook cells other than "# YOUR CODE HERE". We actually highly recommend you do so to gain a better understanding of what is happening. However, during grading, **these changes are ignored**. 
2. You must ensure that all your code for the particular task is available in the cells that say "# YOUR CODE HERE"
3. Every cell that says "# YOUR CODE HERE" is followed by a "raise NotImplementedError". You need to remove that line. During grading, if an error occurs then you will not receive points for your work in that section.
4. If your code passes the "assert" statements, then no output will result. If your code fails the "assert" statements, you will get an "AssertionError". Getting an assertion error means you will not receive points for that particular task.
5. If you edit the "assert" statements to make your code pass, they will still fail when they are graded since the "assert" statements will revert to the original. Make sure you don't edit the assert statements.
6. We may sometimes have "hidden" tests for grading. This means that passing the visible "assert" statements is not sufficient. The "assert" statements are there as a guide but you need to make sure you understand what you're required to do and ensure that you are doing it correctly. Passing the visible tests is necessary but not sufficient to get the grade for that cell.
7. When you are asked to define a function, make sure you **don't** use any variables outside of the parameters passed to the function. You can think of the parameters being passed to the function as a hint. Make sure you're using all of those variables.
8. Finally, **make sure you run "Kernel > Restart and Run All"** and pass all the asserts before submitting. If you don't restart the kernel, there may be some code that you ran and deleted that is still being used and that was why your asserts were passing.

# Getting Started

Usually, getting your environment set up is one of the most time consuming and frustrating parts of development. If you are seeing this running locally on your computer, congratulations!! You're past most of the hardship.

During class, we'll be using [colab](https://g.co/colab) for our exercises but you may wish to [install Python](https://www.python.org/getit/) on your computers, [pipenv](https://docs.pipenv.org/) and set up [Jupyter](http://jupyter.org/). If you need help with this, contact an instructor or TA. 

In this exercise we're going to briefly go over Jupyter Notebooks, Markdown, Python, data visualization code and the machine learning workflow. There's a lot to cover so let's get started.

## Jupyter

Right now, you're looking at a Jupyter notebook. Jupyter notebooks allow you to interactively run code while maintaining variables' states in memory. Notebooks are extremely useful for getting your ideas tested and quickly iterating on code. Once you've verified and validated your ideas, you should place any function definitions and workflow procedures in `.py` files. 

Notebooks have multiple cells, each cell has a code type and contents as well as output or display. This cell you're looking at is a Markdown cell. You can double click it to see (and edit) its contents. 

If you haven't used Markdown before, this [cheatsheet](https://github.com/adam-p/markdown-here/wiki/Markdown-Cheatsheet) is very useful. A portion of its contents are copied in the next cell for you to edit and experiment with.

## Markdown Cheat Sheet

**Double Click Me To Edit!!!**

__If you already know markdown, you can skip this whole cell.__

<-- Cheatsheet contents -->

### Headings

# H1
## H2
### H3
#### H4
##### H5
###### H6

Alternatively, for H1 and H2, an underline-ish style:

Alt-H1
======

Alt-H2
------

### Emphasis

Emphasis, aka italics, with *asterisks* or _underscores_.

Strong emphasis, aka bold, with **asterisks** or __underscores__.

Combined emphasis with **asterisks and _underscores_**.

Strikethrough uses two tildes. ~~Scratch this.~~

### Lists

1. First ordered list item
2. Another item
⋅⋅* Unordered sub-list. 
1. Actual numbers don't matter, just that it's a number
⋅⋅1. Ordered sub-list
4. And another item.

⋅⋅⋅You can have properly indented paragraphs within list items. Notice the blank line above, and the leading spaces (at least one, but we'll use three here to also align the raw Markdown).

⋅⋅⋅To have a line break without a paragraph, you will need to use two trailing spaces.⋅⋅
⋅⋅⋅Note that this line is separate, but within the same paragraph.⋅⋅
⋅⋅⋅(This is contrary to the typical GFM line break behaviour, where trailing spaces are not required.)

* Unordered list can use asterisks
- Or minuses
+ Or pluses

### Links


[I'm an inline-style link](https://www.google.com)

[I'm an inline-style link with title](https://www.google.com "Google's Homepage")


### Images

Here's a logo (hover to see the title text):

Inline-style: 
![alt text](https://github.com/adam-p/markdown-here/raw/master/src/common/images/icon48.png "Logo Title Text 1")

Reference-style: 
![alt text][logo]


### Code

Inline `code` has `back-ticks around` it.

```javascript
var s = "JavaScript syntax highlighting";
alert(s);
```
 
```python
s = "Python syntax highlighting"
print s
```
 
```
No language indicated, so no syntax highlighting. 
But let's throw in a <b>tag</b>.
```

### Tables

Colons can be used to align columns.

| Tables        | Are           | Cool  |
|---------------|---------------|-------|
|1|2|3|
|4|5|6|
|7|8|9|

There must be at least 3 dashes separating each header cell.
The outer pipes (|) are optional, and you don't need to make the 
raw Markdown line up prettily.


Jupyter organizes content in cells. Each cell has a type such as Code or Markdown and can be run to produce an output or render itself. 

Jupyter also has some cool features that allow us to apply commands to a cell. These are called [magic commands](https://ipython.readthedocs.io/en/stable/interactive/magics.html). 

We use magic commands by typing `%` for line magic and `%%` for cell magic. Let's try the `%%time` magic command.

```python
import time
```

```python
%%time
time.sleep(10) # sleep for 10 seconds
```

Output:
```
Wall time: 10 s
```

## Packages / Libraries

You probably noticed above that we ran this line of code `import time` what that's doing is importing the python standard library's [time module](https://docs.python.org/3/library/time.html). Part of why Python is so popular is the availability of awesome libraries we can import. The `time` module comes with your Python install but we're also going to want to use 3rd party libraries. 

Python libraries are generally hosted on [pypi.org](https://pypi.org/) (the Python Package Index) and we can install them using [pip](https://pip.pypa.io/en/stable/) to install them to our global Python installation. For project-specific package management, it is best to setup a virtual environemnt from [pipenv](https://github.com/pypa/pipenv) or at least [virtualenv](https://virtualenv.pypa.io/en/latest/).

## Shell Commands

Now one useful feature in Jupyter notebooks is that we can run shell commands from the notebook. We do that by adding `!` before whatever we're typing.

```python
!echo "Hello!"
```

Output:
```
"Hello!"
```

**However, using shell commands causes Jupyter Notebooks to only work correctly on systems which implement the shell commands referenced**. This becomes a problem when systems use different commands to perform the same actions or when systems do not implement the shell commands required by a notebook.

For example, POSIX compliant systems (Linux, MacOS, Colab) use `pwd` to list files in the current directory while Windows systems use `chdir`. So, only one of the below code cells should execute properly on your system (the other will fail to execute).

```python
# Executes correctly on POSIX compliant systems (Colab, Mac, Linux, etc.)
! pwd
```

Output:
```
'pwd' is not recognized as an internal or external command,
operable program or batch file.
```

```python
# Executes correctly on Windows systems
! chdir
```

Output:
```
C:\Users\chrisbulldog\Desktop\Machine Learning 470\Assignments
```

## Magic functions

Notice that when a shell command fails, it might not return anything. This is because shell commands are executed within a sub-shell and the data in [stdout](https://www.cplusplus.com/reference/cstdio/stdout/) is returned to the primary Jupyter Notebook environment to be used.

[stdout](https://www.cplusplus.com/reference/cstdio/stdout/) does not contain any information about errors (this information is sent to [stderr](https://www.cplusplus.com/reference/cstdio/stderr/) instead). So, when a shell command does not execute correctly, nothing is returned because [stdout](https://www.cplusplus.com/reference/cstdio/stdout/) is empty.

Jupyter Notebooks is built on top of IPython. To improve cross-system compatibility, IPython implements many of the most popular shell commands as magic functions (`%`) which work the same on all systems. A list of these magic functions can be found [here](https://ipython.readthedocs.io/en/stable/interactive/magics.html). **For this reason, it is highly recommended that magic functions `%` be favored over shell commands `!` when possible.**

**Note**: Magic functions (`%`) execute within the primary Jupyter Notebook environment. This is different than shell commands (`!`) which execute within sub-shells.

The following code cell illustrates the differences between the two commands. The demonstration is fairly nuanced, so don't worry if you don't fully understand it. Just recognize that magic commands (`%`) are better suited for use across different operating systems than shell commands (`!`).

```python
# This code cell only works on POSIX compliant systems
# Change this variable to `True` if your system is POSIX compliant (ie. Colab, Linux, Mac, etc.)
is_posix_compliant = False

if is_posix_compliant:
    # Save current directory
    cur_dir = !pwd
    cur_dir = cur_dir[0]

    # Get directory contents
    original_dir_contents = !ls
    
    # Move up one directory within sub-shell
    !cd ..

    # Get directory contents
    sub_shell_dir_contents = !ls
    
    # Move up one directory within Jupyter Notebook environment
    # (Using `-q` = `quiet` so nothing should be printed)
    %cd -q ..
    
    # Get directory contents
    magic_function_dir_contents = !ls
    
    # The directory contents before and after the sub-shell cd are the same
    # But both are different from that of directory moved to by the magic function
    assert original_dir_contents == sub_shell_dir_contents != magic_function_dir_contents
    
    # Send Jupyter Notebook working directory back to original location
    # Note the use of brackets indicating a variable from the Jupyter Notebook environment must be used
    %cd -q {cur_dir}
```

## Pip

To install a package, we simply run `pip install <package-name>`. The main library we'll be using for the first part of the course, is [scikit-learn](https://scikit-learn.org/). You can install it by running `pip install scikit-learn`. 

If you're running this notebook from colab, that lbrary is probably already installed for you. When installing, make sure you're clear on where the package will be installed. Some important questions to ask are:

1. Do I want to install this package locally for my project or globally for every project?
1. Which Python installation is this going to be installed for?
1. Which Python version is that?

```python
# Uncomment the below line to install scikit-learn
# ! pip install scikit-learn
```

## Python

Let's go over some basic Python.

We can do some basic maths as follows:

```python
1 + 1
```

## Cell output
Note that when a jupyter notebook cell is run, if it has multiple lines of code that would otherwise generate printed text output in a terminal, only the output generated by the final code line is displayed below the cell.

```python
x = 5
y = 2
x+y
x/y
```

There are other ways to print out multiple results, as in the two examples below.

```python
x = 5
y = 2
# Do multiple calculations, and put them into a tuple
x + y, x / y, x // y, x **y, x % y
```

```python
# Use the print function
x = 5
y = 2
print(f"The sum is {x+y}.")
print(f"The dividend is {x/y}.")
print(f"The integer dividend is {x//y}.")
print(f"5 squared is {x **y}.")
print(f"The remainder is {x % y}.")
```

We can also do some simple logic flow:

```python
score = 95
if score > 90:
    letterGrade = "A"
elif score > 80:
    letterGrade = "B"
elif score > 70:
    letterGrade = "C"
else:
    letterGrade = "F"

print(f"That score gets a letter grade of {letterGrade}.")
```

This is how we iterate in a loop:

```python
for i in range(1,11):
    if i == 1:
        print("This is the 1st time.")
    elif i == 2:
        print("This is the 2nd time.")
    elif i == 3:
        print("This is the 3rd time")
    else:
        print(f"This is the {i}th time")
```

List comprehensions are a very powerful and useful tool in Python, here are a couple examples:

```python
[i+1 for i in range(10)]
```

```python
[i**2 for i in range(10)]
```

## Passing by assignment

In Python, arguments are passed by assignment.

This means that setting a variable (`a`) equal to another variable (`b`) as shown here:

`a = b`

will assign `a` to a new copy of the reference of the object that `b` is currently referencing

That's a real mouth full so let's look at an example of this and carefully analyze what is happening on each line to see what this actually means:

```python
# Create a string object and reference it using variable `b`
b = "super epic string"

# Create variable `a` and initialize it so that it also references the string that `b` is referencing
a = b

# Create a new string object and reference it using variable `b`
# This causes `b` to abandon its prior reference to "super epic string"
b = "not so epic string"

# `a` still maintains its reference to "super epic string" (as shown by this assertion).
# This is because when `a` was set equal to `b`, it was actually set equal to a new copy of the reference
# to the object that `b` was already referencing
assert a == "super epic string"
assert b == "not so epic string"
assert a != b
```

A string is an immutable type meaning that once it is created, it cannot be changed. A list, on the other hand, is mutable meaning that it can be changed after it is created.

Because variables in Python are references to objects (rather than objects themselves), when two variables reference a mutable type that is changed after its creation, both variables will pick up on the changes.

Let's look at a code snippet to see what this actually means:

```python
# Create a list object and reference it using variable `b`
b = [0, 1, 2, 3]

# Create variable `a` and initialize it so that it also references the list that `b` is referencing
a = b

# Use variable `b` to modify the list it is referencing at the `1th` index
b[1] = 100

# Assert that variable `b` successfully changed the list it is referencing
assert b == [0, 100, 2, 3]

# Assert that variable `a` references the same, changed object
assert a == b == [0, 100, 2, 3]
```

```python
# Now let's create a new list and reference it using variable `b`
b = [0, 1, 2, 3]

# Assert that `a` still references the original list and `b` references the new list
assert a == [0, 100, 2, 3]
assert b == [0, 1, 2, 3]
assert a != b
```

## Copying objects

Sometimes it is necessary to create an entirely new copy of an object rather than a new reference to the same object. For example, it may be necessary to completely duplicate the contents of `b` into `a` so that changes in `b` do not affect `a`.

In these scenarios, copying can be used to create duplicate objects rather than a copy of a reference to the same object. There are two types of copies: shallow copies and deep copies.

A shallow copy is less expensive than a deep copy because it only duplicates objects at the top level whereas a deep copy duplicates all objects at all levels (children + children of children).

In Python, the `.copy()` method is provided with most data types. It offers a simple interface for performing shallow copies.

Let's look at how Python's `.copy()` method actually works!

```python
# Create a list object and reference it using variable `b`
b = [0, 1, 2, 3, ["child_list_0", "child_list_1"]]

# Create variable `a` and initialize it so that it references a duplicate of the list that `b` references
a = b.copy()

# Use variable `b` to modify the list it is referencing at the `1th` index
b[1] = 100

# Assert that variable `b` successfully changed the list it is referencing
assert b == [0, 100, 2, 3, ["child_list_0", "child_list_1"]]

# Assert that variable `a` did NOT register these changes
assert a == [0, 1, 2, 3, ["child_list_0", "child_list_1"]] != b
```

```python
# Now let's modify the inner child list both variables `a` and `b` contain
b[4][0] = "new_child_list_0"

# Assert that variable `b` successfully changed the inner list it is referencing
assert b == [0, 100, 2, 3, ["new_child_list_0", "child_list_1"]]

# Ok, here comes the tricky part
# Variable `a`'s 4th term is a reference to a list (because, as discussed above, Python passes by assignment)
# So, when .copy() was called, a duplicate of the reference to the list OBJECT was created.
# Remember, Python's .copy() is a shallow copy so the inner list object was not also copied.
# Therefore, variable `a`'s 4th index references the same list OBJECT that variable `b`'s 4th index references
# For this reason, we assert that variable `a` registered the changes to the inner list made by variable `b`
assert a == [0, 1, 2, 3, ["new_child_list_0", "child_list_1"]]
```

If you need to create a completely new copy of an object, such that all contents of `a`, including children, are true duplicates of the contents of `b`, one helpful way is to use the `.deepcopy()` function found in the `copy` package.

```python
# Use `pip` to install copy (if you don't already have it)
# ! pip install copy
import copy

b = [0, 1, 2, 3, ["child_list_0", "child_list_1"]]

# Make a deep copy of b
a = copy.deepcopy(b)

# Alter a child (deep) element of b
b[4][0] = "new_child_list_0"

# Assert that variable `b` successfully changed the list it is referencing
assert b == [0, 1, 2, 3, ["new_child_list_0", "child_list_1"]]

# Assert that variable `a` did NOT register these changes
assert a == [0, 1, 2, 3, ["child_list_0", "child_list_1"]] != b
```

## Functions
We can also define some functions in Python. Here's an example function that gets the square root of sum of the squares in a list. Does this sound familiar?

```python
def get_metric(vector):
    """Get the square root of the sum of the squares in an iterable.

    Args:
        vector (iterable): the data to calculate the metric of
    
    Returns:
        float
    """
    return sum([i ** 2 for i in vector])**(1/2)
```

```python
get_metric([3,4])
```

That metric was the $\ell_2$-norm. Can you write a function that takes a vector and value for p and determines the $\ell_p$-norm of a vector? 

<!-- grade_id: get_lp, points: 0, grade: false, solution: true -->
```python
def get_lp(vector, p):
    """Get the $\ell_p$-norm of data

    Args:
        vector (iterable): the data to calculate the metric of
        p (integer): the value of p in the $\
    
    Returns:
        float
    """
    # YOUR CODE HERE
    raise NotImplementedError()
    # NOTE: you will be seeing a lot of these NotImplementedError()'s throughout the semester.
    # You need to take them out once you start working on the specific section of the assignment.     
```

```python
# Use `pip` to install pytest (if you don't already have it)
# ! pip install pytest
import pytest
```

<!-- grade_id: test_get_lp, points: 10, grade: true, solution: false -->
```python
assert get_lp([3,4],2) == 5
assert get_lp([3,4,5],2) == pytest.approx(7,0.1)
assert get_lp([10], 5) == pytest.approx(10)
```

We can also write functions in shorthand using `lambda` this creates an anonymous function. This is useful when we're using maps and filters.

```python
lambda x : x**2
```

```python
example_map = map(lambda x: x**3,[1,2,3,4,5])
list(example_map)
```

```python
example_filter = filter(lambda x: x%2 == 1, [1,2,3,4,5])
list(example_filter)
```

Python has a great ecosystem of libraries available and we will be using them for almost all the code we write. Here's how to import code in Python. I recommend reading up on the intricacies of imports. [1](https://stackoverflow.com/a/710603) [2](https://stackoverflow.com/a/19185936)

You're probably going to see the below imports and couple lines of code very frequently at the top of most files we write. First we import the libraries we often use; [numpy](https://www.numpy.org/), [pandas](https://pandas.pydata.org/), [matplotlib](https://matplotlib.org/), and [scikit-learn](https://scikit-learn.org/). The last two lines are for setting plots to show in the notebook directly and not open up in another window and to select the styleguide for the plots we make to ggplot. You can select another style from the [available options](https://matplotlib.org/gallery/style_sheets/style_sheets_reference.html) if you like.


```python
import numpy as np
import pandas as pd
import matplotlib
import matplotlib.pyplot as plt
import sklearn as sk

%matplotlib inline
plt.style.use("ggplot")
```

## Data Visualization

Let's go over some data visualization basics. To start, we'll need some data. We'll create data from a function with some noise and see if we can fit a polynomial to it. Parts of the following are extracted from the following sklearn tutorials. [1](http://scikit-learn.org/stable/auto_examples/model_selection/plot_underfitting_overfitting.html) [2](http://scikit-learn.org/stable/auto_examples/model_selection/plot_validation_curve.html)

```python
def underlying_pattern(X, noise=0.1):
    """Create a function to fit
    
    Args:
        X numpy.array : array of x values
        
    Returns:
        numpy.array : array of output values
    """
    return np.cos(1.5 * np.pi * X) + np.random.randn(X.shape[0]) * noise
```

You should always set a seed value so that you get reproducable results.

```python
np.random.seed(0)
```

```python
n = 200
X = np.sort(np.random.rand(n))
y = underlying_pattern(X)
```

```python
plt.plot(X, y, marker="o", linestyle="None")
# plt.scatter(X, y) # This might be easier to remember, both functions accomplish the same thing. Try switching to the commented one.
plt.xlabel("x")
plt.ylabel("y")
plt.title("Sample Data")
plt.xlim((0, 1))
plt.ylim((-2, 2))
plt.show()
```

## Machine Learning Workflow

Now that we have some data, let's try to fit a machine learning model, the one we already covered, polynomial regression, to the data. For now, you'll get all the code you need to run this experiment.

```python
from sklearn.model_selection import train_test_split
from sklearn.pipeline import Pipeline
from sklearn.preprocessing import PolynomialFeatures
from sklearn.linear_model import LinearRegression
```

```python
X_train, X_test, y_train, y_test = train_test_split(X, y, random_state=0, test_size=0.2)
X_model, X_valid, y_model, y_valid = train_test_split(X_train, y_train, random_state=0, test_size=0.2)

print(f"All Data:        {len(X)} points")
print(f"Training data:   {len(X_train)} points")
print(f"Testing data:    {len(X_test)} points")
print(f"Modeling data:   {len(X_model)} points")
print(f"Validation data: {len(X_valid)} points")
```

```python
plt.plot(X_model, y_model, marker="o", linestyle="None")
# plt.scatter(X, y) # This might be easier to remember, both functions accomplish the same thing. Try switching the commented one.
plt.xlabel("x")
plt.ylabel("y")
plt.title("Model Data")
plt.xlim((0, 1))
plt.ylim((-2, 2))
plt.show()
```

Here you can train a model by manually adjusting the hyperparameter, in this case the polynomial order.

```python
import math
```

```python
## Change this value and check out the resulting plot
order = 1


## Don't worry about the below code for now
xMin = math.floor(X_model.min())
xMax = math.ceil(X_model.min())
polynomial_features = PolynomialFeatures(degree=order, include_bias=False)
linear_regression = LinearRegression()
pipeline = Pipeline([("polynomial_features", polynomial_features),
                         ("linear_regression", linear_regression)])
pipeline.fit(X_model[:, np.newaxis], y_model)

plt.plot(X_model, y_model, marker="o", linestyle="None")
plt.plot(np.linspace(xMin,xMax), pipeline.predict(np.linspace(xMin,xMax)[:, np.newaxis]))
plt.xlabel("x")
plt.ylabel("y")
plt.title(f"Order {order}")
plt.xlim((0, 1))
plt.ylim((-2, 2))
plt.tight_layout()
plt.show()
```

By switching the order we can see the new models generated. At first, using a larger order does better but after we reach a certain point, the model begins to overfit. Let's try to determine the best value for the hyperparemeter. In order to do that, we must select a metric to use to measure the effectiveness of a model. You can select any of the available [regression metrics from sklearn](http://scikit-learn.org/stable/modules/classes.html#regression-metrics) or develop your own.

In this example, we will start with the mean squared error. A better model will have a lower value for mean squared error.

```python
from sklearn.metrics import mean_squared_error
```

```python
# Feel free to change up to which order you go up to
orders = range(1,50)

scores = []
training_scores = []
for order in orders:
    polynomial_features = PolynomialFeatures(degree=order, include_bias=False)
    linear_regression = LinearRegression()
    pipeline = Pipeline([("polynomial_features", polynomial_features),
                             ("linear_regression", linear_regression)])
    pipeline.fit(X_model[:, np.newaxis], y_model)
    training_scores.append(mean_squared_error(y_model, pipeline.predict(X_model[:, np.newaxis])))
    scores.append(mean_squared_error(y_valid, pipeline.predict(X_valid[:, np.newaxis])))
```

```python
plt.plot(orders, training_scores, label="Training Score")
plt.plot(orders, scores, label="Validation Score")
plt.title("Model Comparison")
plt.xlabel("Order")
plt.ylabel("Mean Squared Error")
plt.legend(loc="best")
plt.show()
```

```python
# Feel free to change the count of values to plot
plot_first = 20
plt.plot(orders[:plot_first], training_scores[:plot_first], label="Training Score")
plt.plot(orders[:plot_first], scores[:plot_first], label="Validation Score")
plt.title("Model Comparison")
plt.xlabel("Order")
plt.ylabel("Mean Squared Error")
plt.legend(loc="best")
plt.show()
```

Now based on the above figure, select an order and let's see how it performs on the test set.

<!-- grade_id: selected_order, points: 0, grade: false, solution: true -->
```python
# Set the value of selectedOrder to what you think the best hyperparameter value should be
# YOUR CODE HERE
raise NotImplementedError()
```

```python
polynomial_features = PolynomialFeatures(degree=selectedOrder, include_bias=False)
linear_regression = LinearRegression()
pipeline = Pipeline([("polynomial_features", polynomial_features),
                         ("linear_regression", linear_regression)])

# Now we train on all the training data
pipeline.fit(X_train[:, np.newaxis], y_train)
# And test on the test data
test_score = mean_squared_error(y_test, pipeline.predict(X_test[:, np.newaxis]))
train_score = scores[selectedOrder-1]

print(f"The selected hyperparameter obtained a score of {train_score} in validation and a score of {test_score} in testing.")
if abs(test_score - train_score) < 0.01 or test_score < train_score:
    print("Good job!")
else:
    print("You might want to try again!")
```

<!-- grade_id: test_selected_order, points: 10, grade: true, solution: false -->
```python
assert abs(test_score - train_score) < 0.01 or test_score < train_score
```

## Feedback

<!-- grade_id: feedback, points: 0, grade: false, solution: true -->
```python
def feedback():
    """Provide feedback on the contents of this exercise
    
    Returns:
        string
    """
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: test_feedback, points: 0, grade: true, solution: false -->
```python

```